---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`
Create a variable a, equal to 5 and print out a

`@hint`


`@pre_exercise_code`
```{r}
b <- 5
```

`@sample_code`
```{r}

```

`@solution`
```{r}
a <- 5
a
```

`@sct`
```{r}

```
